Uncovering the electrical synapse proteome in retinal neurons via in vivo proximity labeling

Author response:

The following is the authors’ response to the previous reviews

Reviewer 1

The authors should clarify the statement regarding the expression in horizontal cells (lines 170-172). In line 170, it is stated that GFP was observed in horizontal cells. Since GFP is fused to Cx36, the observation of GFP in horizontal cells would suggest the expression of Cx36-GFP.

We believe that there appears to be a misunderstanding. GFP is observed in horizontal cells, because the test AAV construct, which consists of the HKamac promoter and a downstream GFP sequence, was used to validate the promoter specificity in wildtype animals. This was just a test to confirm that HKamac is indeed active in AII amacrine cells as previously described by Khabou et al. 2023. This construct was not used for the large scale BioID screen. For these experiments, V5-dGBP-Turbo was expressed under the control of the HKamac promoter as illustrated in Figure 2A.

Fig 7: the legend is missing the descriptions for panels A-C.

We apologize for this mistake. We have missed the label “(A-C)” and added it to the legend.

Supplemental files are not referenced in the manuscript.

We have added a reference for these files in line 221-226.

Reviewer 2

Supplementary Files 1 and 2 are presented as two replicates of the zebrafish proteomic datasets, but they appear to be identical.

This appears to be a misunderstanding. These two replicates contain slightly different hits, although the most abundant candidates are identical.

Reviewer 3

Thank you for the positive comments

Changes from the earlier version to the revised version

Removed in the revised version: Abstract

Abstract

Electrical synapses containing Connexin 36 (Cx36) represent the main means for direct electrical communication among neurons in the mammalian nervous system. However, little is known about the protein complexes that constitute these synapses. In the present study, we applied different BioID strategies to screen the interactomes of Connexin 36 and its zebrafish orthologue Cx35b in retinal neurons. For in vivo proximity labeling in mice, we took advantage of the Cx36-EGFP strain and expressed a GFP-nanobody-TurboID fusion construct selectively in AII amacrine cells. For in vivo BioID in zebrafish, we generated a transgenic line expressing a Cx35b-TurboID fusion under control of the Cx35b promoter. Both strategies allowed us to capture a plethora of molecules that were associated with electrical synapses and showed a high degree of evolutionary conservation in the proteomes of both species. Besides known interactors of Cx36 such as ZO-1 and ZO-2 we have identified more than 50 new proteins, such as scaffold proteins, adhesion molecules and regulators of the cytoskeleton. Moreover, we determined the subcellular localization of these proteins in mouse retina and tested potential binding interactions with Cx36. Amongst these new interactors, we identified signal induced proliferation associated 1 like 3 (Sipa1l3), a protein that has been implicated in cell junction formation and cell polarity, as a new scaffold of electrical synapses. Interestingly, Sipa1l3 was able to interact with ZO-1, ZO-2 and Cx36, suggesting a pivotal role in electrical synapse function. In summary, our study provides the first detailed view of the electrical synapse proteome in retinal neurons, which is likely to apply to electrical synapses elsewhere.

Introduction

With the discovery of acetylcholine as the first neurotransmitter at the beginning of the 20th century, it was widely accepted that chemical synaptic transmission is the only means of communication for neurons. Decades later, several authors independently demonstrated evidence for the existence of electrical synapses in invertebrates as well as in vertebrates, and subsequent studies identified gap junctions, which are composed of clusters of intercellular channels, as the morphological substrates of these synapses (1, 2). Today it is considered common knowledge among neuroscientists that both types of synapses coexist in the nervous system and that each of them provides unique features essential for accurate signal processing. Yet, compared to the efforts that have been made to study protein complexes underlying chemical synapse function and regulation, we know fairly little about the proteome that constitutes an electrical synapse (3).

In contrast to chemical synapses, electrical synapses don’t rely on neurotransmitters and specific receptor molecules but instead make use of intercellular channels to convey neural signals bidirectionally in the form of ionic currents (4). This seemingly simple mode of communication allows for an instantaneous signal transmission that synchronizes the activities of electrically coupled neurons giving rise to complex brain functions such as network oscillation, hormone secretion or coordination of motor functions (5, 6).

Earlier version: In the mammalian nervous system, electrical synapses are mainly, but not exclusively formed by Connexin 36 (Cx36), a connexin isoform that is highly specific for neurons (7–9) and often referred to as the “major neuronal connexin”. One exception to the rule are the insulin producing beta cells in the pancreas (10). In general, Cx36 gap junction channels are known to exhibit a very low single channel conductance (11) but a remarkable degree of plasticity that allows neurons to regulate the number of open channels and hence the strength of synaptic transmission in order to adapt to environmental stimuli (12). A prime example of a neuron that is heavily coupled via Cx36 and in which the signal transduction pathways that regulate electrical coupling have been well described is the AII amacrine cell (13–15). The AII cell is a glycinergic interneuron in the mammalian retina that is electrically coupled to neighboring AII amacrine cells and to ON Cone bipolar cells (16). The main task of this cell is to collect neural signals that originate in rods and relay them into the cone pathway. This circuit is referred to as the ‘primary rod pathway’ and it forms the neurophysiological basis for scotopic (dim light) vision. Within the primary rod pathway, electrical coupling of AII amacrine cells is regulated in an activity dependent manner (13). The opening of Cx36 channels in AII amacrine cells requires the activation of calcium calmodulin dependent kinase II (CaMKII), which is triggered by the influx of calcium through non-synaptic NMDA receptors. Once CaMKII is activated, it binds to Cx36 and phosphorylates the intracellular domains causing the channel to open (13, 17, 18). These activity driven changes clearly resemble long-term potentiation as it has been described for chemical synapses in hippocampal neurons (19). Another, almost obvious, but not less important, parallel between electrical and chemical synapses, is the their basic “architecture”: 1) Glutamate receptors, for instance, interact with PSD95 (20), a membrane associated guanylate kinase (MAGUK) that is similar in function and domain structure to ZO-1, another MAGUK and well-known interactor of Cx36. 2) Electrical synapses are surrounded by adhesions (21). 3) As described for glutamatergic synapses, electrical synapses also harbor electron dense structures that appear to consist of a dense assemblage of synaptic proteins (22, 23). However, the composition of this density or the type of adhesion molecule that is necessary to specify where an electrical synapse is formed still remain a mystery and require sophisticated proteomics to solve.

Revised version: In the mammalian nervous system, electrical synapses are mainly, but not exclusively formed by Connexin 36 (Cx36), a connexin isoform that is highly specific for neurons (7–9) and often referred to as the “major neuronal connexin”. One exception to the rule are the insulin producing beta cells in the pancreas (10). In general, Cx36 gap junction channels are known to exhibit a very low single channel conductance (11) but a remarkable degree of plasticity that allows neurons to regulate the number of open channels and hence the strength of synaptic transmission in order to adapt to environmental stimuli (12). A prime example of a neuron that is heavily coupled via Cx36 and in which the signal transduction pathways that regulate electrical coupling have been well described is the AII amacrine cell (13–15). The AII cell is a glycinergic interneuron in the mammalian retina that is electrically coupled to neighboring AII amacrine cells and to ON Cone bipolar cells (16). The main task of this cell is to collect neural signals that originate in rods and relay them into the cone pathway. This circuit is referred to as the primary rod pathway and it forms the neurophysiological basis for scotopic vision. Within the primary rod pathway, electrical coupling of AII amacrine cells is regulated in an activity dependent manner (13). The opening of Cx36 channels in AII amacrine cells requires the activation of calcium calmodulin dependent kinase II (CaMKII), which is triggered by the influx of calcium through non-synaptic NMDA receptors. Once CaMKII is activated, it binds to Cx36 and phosphorylates the intracellular domains causing the channel to open (13, 17, 18). These activity driven changes clearly resemble long-term potentiation as it has been described for chemical synapses in hippocampal neurons (19). Another, almost obvious, but not less important, parallel between electrical and chemical synapses, is the their basic “architecture”: 1) Glutamate receptors, for instance, interact with PSD95 (20), a membrane associated guanylate kinase (MAGUK) that is similar in function and domain structure to ZO-1, another MAGUK and well-known interactor of Cx36. 2) Electrical synapses are surrounded by adhesions (21). 3) As described for glutamatergic synapses, electrical synapses also harbor electron dense structures that appear to consist of a dense assemblage of synaptic proteins (22, 23). However, the composition of this density or the type of adhesion molecule that is necessary to specify where an electrical synapse is formed still remain a mystery and require sophisticated proteomics to solve.

The task of identifying compartment specific proteomes or interactomes for a given protein of interest (POI) with standard techniques has always been challenging due to the abundance of non-specific background and the requirement to preserve intact protein complexes. To overcome these limitations, several laboratories have developed novel proximity techniques such as TurboID (24, 25), which utilize promiscuous biotin ligases to biotinylate neighboring proteins within a radius of 5 nm of the POI. Like any other tag, TurboID can be fused to the POI and is then expressed in cells, where it is allowed to biotinylate proximal proteins that are subsequently isolated with streptavidin beads under stringent buffer conditions. The ability to directly label “associated” proteins with an affinity tag combined with the fast enzyme kinetics have made TurboID a valuable tool for modern in vivo proteomics. Previously, BioID and TurboID have been used to study the proteomes of chemical synapses and even contact sites of astrocytes and neurons in nematodes and mice (26–28).

In the present study, we have applied two different TurboID approaches to uncover the electrical synapse proteome in the mouse and the zebrafish retina. Contrasting these two model organisms allowed us to expose essential components of the molecular architecture of vertebrate electrical synapses. Beyond known components of electrical synapses, such as ZO-1 and ZO-2, our screen identified a plethora of new molecules that were not previously shown to associate with Cx36/Cx35b. These include components of the endocytic machinery, adhesion molecules, regulators of the cytoskeleton, membrane trafficking, as well chemical synapse proteins. Overall, our screen showed that many components of electrical synapses in zebrafish and mouse retinas were highly conserved, in particular the ZO proteins and the endocytosis components, indicating a high degree of evolutionary conservation that suggests a critical functional role for these proteins.

Results

Two TurboID approaches allow a glimpse at the electrical synapse proteome in zebrafish photoreceptors and AII amacrine cells of the mouse retina

Earlier version: In this study, we have adopted two in vivo BioID approaches to gain a more comprehensive understanding of protein complexes that are associated with the neuronal connexin Cx36 and its zebrafish orthologue Cx35b in retinal neurons. To implement BioID in zebrafish, we generated a transgenic line expressing a Cx35b-V5-TurboID fusion protein (we used an internal insertion site as previously described (29)) under the control of Cx35b promoter (Figure 1 A). As shown in control experiments, the expression of this construct recapitulated the expression pattern of wildtype Cx35b in photoreceptors (30) which were labelled with ZPR1, indicating that Cx35b-V5-TurboID is correctly targeted to synapses (Figure 1B). The inner retina, however, displayed much weaker expression of the fusion construct, which is why we anticipated that most of the candidates we detected with this strategy are proteins that are associated with Cx35b in photoreceptors. To induce efficient biotinylation in zebrafish we injected 10mM biotin solution intraperitoneally on three consecutive days and isolated biotinylated proteins for mass spectrometry (Figure 1B and C). For the analysis of mass spec data sets we compared the abundance of candidate proteins from the Cx35b-V5-TurboID strain to wildtype samples that did not express TurboID. In Cx35b-V5-TurboID samples, but not in wild type retinas, we found several known interactors of the Gjd2 family such as scaffold proteins including Tjp1a, Tjp1b, Tjp2a, Tjp2b, Mupp1, Psd95 and Af6 (Figure 1D) (31–34). In addition to these candidates, we identified proteins that are involved in membrane trafficking including Syt4, Sec22B, components of the endosomal trafficking machinery Snap91 and Eps15I1a, and an additional Cx36 orthologue termed Cx34.7, which has previously been localized to photoreceptors in bass retina (35). The detection of this Gjd1 isoform via BioID was somewhat surprising, as Cx34.7 was shown to form gap junctions that are separated from Cx35b containing clusters in bass photoreceptors. To determine the subcellular localization of Cx34.7 and its connection to Cx35b in zebrafish photoreceptors, we generated a transgenic Cx34.7-GFP strain that expressed an internally GFP-tagged fusion protein under the control of the Cx34.7 promoter. In this strain, we found that Cx34.7 was only expressed in the outer plexiform layer (OPL) where it was associated with Cx35b at some gap junctions (Figure 1E). Thus, besides known interactors of the Gjd family, our in vivo BioID approach identified an additional connexin isoform that is connected to Cx35b in photoreceptors.

Revised version: In this study, we have adopted two in vivo BioID approaches to gain a more comprehensive understanding of protein complexes that are associated with the neuronal connexin Cx36 and its zebrafish orthologue Cx35b in retinal neurons. To implement BioID in zebrafish, we generated a transgenic line expressing a Cx35b-V5-TurboID fusion protein (we used an internal insertion site as previously described (29)) under the control of Cx35b promoter (Figure 1 A). As shown in control experiments, the expression of this construct recapitulated the expression pattern of wildtype Cx35b in photoreceptors (30) which were labelled with ZPR1, indicating that Cx35b-V5-TurboID is correctly targeted to synapses (Figure 1B). The inner retina, however, displayed much weaker expression of the fusion construct, which is why we anticipated that most of the candidates we detected with this strategy are proteins that are associated with Cx35b in photoreceptors. To induce efficient biotinylation in zebrafish we injected 10mM biotin solution intraperitoneally on three consecutive days and isolated biotinylated proteins for mass spectrometry (Figure 1B and C). A complete list of proteins captured in two experiments is shown in Supplement 1. For the analysis of mass spec data sets we compared the abundance of candidate proteins from the Cx35b-V5-TurboID strain to wildtype samples that did not express TurboID, and evaluated specificity using Significance Analysis of Interactome (SAINT) (31). In Cx35b-V5-TurboID samples, but not in wild type retinas, we found several known interactors of the Gjd2 family such as scaffold proteins including Tjp1a, Tjp1b, Tjp2a, Tjp2b, Mupp1, Psd95 and Af6 (Figure 1D) (32–35). In addition to these candidates, we identified proteins that are involved in membrane trafficking including Syt4, Sec22b, components of the endosomal trafficking machinery Snap91 and Eps15I1a, and an additional Cx36 orthologue termed Cx34.7, which has previously been localized to photoreceptors in bass retina (36). The detection of this Gjd1 isoform via BioID was somewhat surprising, as Cx34.7 was shown to form gap junctions that are separated from Cx35b containing clusters in bass photoreceptors. To determine the subcellular localization of Cx34.7 and its connection to Cx35b in zebrafish photoreceptors, we generated a transgenic Cx34.7-GFP strain, that expressed an internally GFP-tagged fusion protein under the control of the Cx34.7 promoter. In this strain, we found that Cx34.7 was only expressed in the outer plexiform layer (OPL) where it was associated with Cx35b at some gap junctions (Figure 1E). Thus, besides known interactors of the Gjd family, our in vivo BioID approach identified an additional connexin isoform that is connected to Cx35b in photoreceptors.

Generation of a Cx35b-V5-TurboID zebrafish line.

Earlier version: (A) Cartoon illustrating the generation of Cx35-V5-TurboID fish via Tol2-mediated transgenesis and outline of in vivo biotinylation experiments. To induce efficient proximity labeling, zebrafish were intraperitoneally injected with 30µl of 5mM biotin (PBS) for three consecutive days. Afterwards the animals were sacrificed, and the retinas were isolated for streptavidin pull-downs. (B) Confocal scans of entire zebrafish retinas, the outer retina including ZPR1 labeled cone photoreceptors to confirm successful targeting of Cx35b-V5-TurboID to photoreceptor gap junctions. Neutravidin Oregon Green labeling was used to validate efficient proximity biotinylation of Cx35b and surrounding molecules in biotin injected Cx35b-V5-TurboID fish. Reagent or antibody used for labeling is shown in the gray box above each image. Scale 10µm. (C) Western blot of streptavidin pull-downs probed with V5 antibodies and streptavidin-HRP. The arrows indicate the position of the Cx35-V5-TurboID construct, which is detected with streptavidin-HRP and a V5 antibody. (D) String diagram illustrating the protein-protein network surrounding Cx35b. Proteins that were two times or more abundant in comparison to the control condition were included in the string diagram. (E) Colocalization of Cx35b and Cx34.7-GFP in the outer plexiform layer of the zebrafish retina suggests that a subset of photoreceptor gap junctions contain both connexins. Scale: 5 µm; inset scale: 0.5 µm.

Revised version: (A) Cartoon illustrating the generation of Cx35-V5-TurboID fish via Tol2-mediated transgenesis and outline of in vivo biotinylation experiments. To induce efficient proximity labeling, zebrafish were intraperitoneally injected with 30µl of 5mM biotin (PBS) for three consecutive days. Afterwards the animals were sacrificed, and the retinas were isolated for streptavidin pull-downs. (B) Confocal scans of entire zebrafish retinas, the outer retina including ZPR1 labeled photoreceptors to confirm successful targeting of Cx35b-V5-TurboID to photoreceptor gap junctions. Neutravidin Oregon Green labeling was used to validate efficient proximity biotinylation of Cx35b and surrounding molecules in biotin injected Cx35b-V5-TurboID fish. Reagent or antibody used for labeling is shown in the gray box above each image. Scale 10µm. (C) Western blot of streptavidin pull-downs probed with V5 antibodies and streptavidin-HRP. The arrows indicate the position of the Cx35-V5-TurboID construct, which is detected with streptavidin-HRP and a V5 antibody. (D) String diagram illustrating the protein-protein network surrounding Cx35b. Proteins that were two times or more abundant in comparison to the control condition were included in the string diagram. (E) Colocalization of Cx35b and Cx34.7-GFP in the outer plexiform layer of the zebrafish retina suggests that a subset of photoreceptor gap junctions contain both connexins. Scale: 5 µm; inset scale: 0.5 µm.

Earlier version: To unravel the Cx36 interactome in the mouse retina we initially delivered a Cx36-V5-TurboID fusion via intravitreal AAV injections into retinal neurons. This construct was similarly designed to the Cx35b clone but regulated by a short human synapsin promoter to drive expression in the neural retina (Figure 2A). Unfortunately, this strategy produced massive overexpression artefacts and resulted in an accumulation of Cx36-V5-TurboID in ganglion cell somas, which would have made an interactome analysis of a synaptic protein impossible. To solve this issue and to ensure that TurboID is targeting electrical synapses, we took advantage of the GFP-directed TurboID approach (36) and the Cx36-EGFP transgenic mouse strain (37–39). The GFP-directed TurboID strategy was initially developed as a modular system for zebrafish by Xiong et al., and utilizes and a destabilized GFP nanobody that is fused to TurboID (36) to direct the biotin ligase towards a given protein of interest carrying a GFP-Tag. To further increase the specificity of our approach and identify the Cx36 interactome within a single type of inhibitory interneuron, the AII amacrine cell, we expressed V5-TurboID-dGBP under control of the HKamac promoter (40). As shown previously, in retinas that were infected with AAVs carrying an HKamac GFP vector, we observed GFP expression mainly in AII amacrine cells but not in bipolar cells (Figure 2B, lack of GFP expression in SCGN labelled neurons). In our hands, we also observed GFP labeling in horizontal cells, which was neglectable for our paradigm, since these neurons do not express Cx36 (41) and because dGBP is degraded unless it is bound to GFP (42). In an initial experiment we injected AAVs carrying HKamac_V5-TurboID-dGBP intravitreally into Cx36-EGFP mice and tracked the localization of the construct two weeks post injection. As expected, all V5 labeled puncta colocalized with Cx36-EGFP, indicating that V5-TurboID-dGBP reached electrical synapses in AII amacrine cells. In the next experiment, we injected a 5 mM biotin solution subcutaneously on 4 consecutive days, 3 1/2 weeks post infection. This treatment was sufficient to induce efficient biotinylation (Figure 1B, lower panel and C), and allowed us to capture a plethora of molecules that were associated with Cx36. To distinguish these proteins from background, we included wild type mice injected with AAV HKamac_V5-TurboID-dGBP throughout the entire experiment and compared the abundance of all proteins we detected. A mass spec hit that was three times or even more abundant in the Cx36-EGFP condition was considered a candidate that is likely to be associated with Cx36. We illustrated the relationship of all proteins that fell into this category in a string diagram (Figure 2D). Among the most abundant proteins in our screen, we identified the ZO proteins ZO-1 and ZO-2 and signal induced proliferation associated 1 like 3 (Sipa1l3), a PDZ domain containing protein that has been implicated in the regulation of cell adhesion and cell polarity. In addition to these candidates, we also identified proteins that are involved in endocytosis, membrane trafficking, regulation of the cytoskeleton, cell adhesion and chemical synapses. All of these hits covered different functional categories of proteins that have been related to connexins (43–48) and thus appear to represent an authentic interactome for a neuronal gap junction protein.

Revised version: To unravel the Cx36 interactome in the mouse retina we initially delivered a Cx36-V5-TurboID fusion via intravitreal AAV injections into retinal neurons. This construct was similarly designed to the Cx35b clone but regulated by a short human synapsin promoter to drive expression in the neural retina (Figure 2A). Unfortunately, this strategy produced massive overexpression artefacts and resulted in an accumulation of Cx36-V5-TurboID in ganglion cell somas, which would have made an interactome analysis of a synaptic protein impossible. To solve this issue and to ensure that TurboID is targeting electrical synapses, we took advantage of the GFP-directed TurboID approach (37) and the Cx36-EGFP transgenic mouse strain (38–40). The GFP-directed TurboID strategy was initially developed as a modular system for zebrafish by Xiong et al., and utilizes and a destabilized GFP nanobody that is fused to TurboID (37) to direct the biotin ligase towards a given protein of interest carrying a GFP-Tag. To further increase the specificity of our approach and identify the Cx36 interactome within a single type of inhibitory interneuron, the AII amacrine cell, we expressed V5-TurboID-dGBP under control of the HKamac promoter (41). As shown previously, in retinas that were infected with AAVs carrying an HKamac GFP vector, we observed GFP expression mainly in AII amacrine cells but not in bipolar cells (Figure 2B, lack of GFP expression in SCGN labelled neurons). In our hands, we also observed GFP labeling in horizontal cells, which was neglectable for our paradigm, since these neurons do not express Cx36 (42) and because dGBP is degraded unless it is bound to GFP (43). In an initial experiment we injected AAVs carrying HKamac_V5-TurboID-dGBP intravitreally into Cx36-EGFP mice and tracked the localization of the construct two weeks post injection. As expected, all V5 labeled puncta colocalized with Cx36-EGFP, indicating that V5-TurboID-dGBP reached electrical synapses in AII amacrine cells. In the next experiment, we injected a 5 mM biotin solution subcutaneously on 4 consecutive days, 3 1/2 weeks post infection. This treatment was sufficient to induce efficient biotinylation (Figure 2B, lower panel and C) and allowed us to capture a plethora of molecules that were associated with Cx36. To distinguish these proteins from background, we included wild type mice injected with AAV HKamac_V5-TurboID-dGBP throughout the entire experiment. A complete list of proteins captured in two experiments is shown in Supplement 2. We compared the abundance of all proteins we detected and evaluated specificity using SAINT. A mass spec hit that was three times or even more abundant in the Cx36-EGFP condition and with a SAINT score greater than 0.5 was considered a candidate that is likely to be associated with Cx36. We illustrated the relationship of all proteins that fell into this category in a string diagram (Figure 2D). Among the most abundant proteins in our screen, we identified the ZO proteins ZO-1 and ZO-2 and signal induced proliferation associated 1 like 3 (Sipa1l3), a PDZ domain containing protein that has been implicated in the regulation of cell adhesion and cell polarity. In addition to these candidates, we also identified proteins that are involved in endocytosis, membrane trafficking, regulation of the cytoskeleton, cell adhesion and chemical synapses. All of these hits covered different functional categories of proteins that have been related to connexins (44–49) and thus appear to represent an authentic interactome for a neuronal gap junction protein.

A GFP-directed TurboID approach uncovers the Cx36 interactome in AII amacrine cells.

Earlier version: (A) Cartoon illustrating different in vivo BioID approaches used to target Cx36. In strategy 1 we expressed an AAV Cx36-V5-TurboID construct with an internal insertion site under the control of the human synapsin promoter. This strategy greatly resulted in overexpression artefacts causing aberrant protein localization. In strategy 2 we adapted the GFP-directed TurboID strategy developed by Xiong et al., (36) to shuttle a V5-dGBP-TurboID construct to Cx36-EGFP containing gap junctions in retinas of the Cx36-EGFP transgenic strain. (B) Vertical sections of AAV infected retinas confirming cell type-specificity of our AAV vectors. The HKamac promoter developed by Khabou et al. (40) is mainly active in AII amacrine cells and does not show any expression in bipolar cells (SCGN, magenta). V5-dGBP-TurboID colocalizes with Cx36-EGFP in AAV transduced retinas, confirming the successful delivery of TurboID to the gap junction. Scale: 5µm. (C) Western blot confirming successful biotinylation and capture of Cx36-EGFP, biotinylated proteins and V5-TurboID-dGBP. (D) String diagram illustrating the protein-protein interaction network associated with Cx36 in AII amacrine cells. Proteins that were three times or more abundant compared to the control condition were included in the string diagram.

Revised version: (A) Cartoon illustrating different in vivo BioID approaches used to target Cx36. In strategy 1 we expressed an AAV Cx36-V5-TurboID construct with an internal insertion site under the control of the human synapsin promoter. This strategy greatly resulted in overexpression artefacts causing aberrant protein localization. In strategy 2 we adapted the GFP-directed TurboID strategy developed by Xiong et al., (37) to shuttle a V5-dGBP-TurboID construct to Cx36-EGFP containing gap junctions in retinas of the Cx36-EGFP transgenic strain. (B) Vertical sections of AAV infected retinas confirming cell type-specificity of our AAV vectors. The HKamac promoter developed by Khabou et al. (41) is mainly active in AII amacrine cells, and does not show any expression in bipolar cells (SCGN, magenta). V5-dGBP-TurboID colocalizes with Cx36-EGFP in AAV transduced retinas, confirming the successful delivery of TurboID to the gap junction. Scale: 5µm. (C) Western blot confirming successful biotinylation and capture of Cx36-EGFP, biotinylated proteins and V5-TurboID-dGBP. (D) String diagram illustrating the protein-protein interaction network associated with Cx36 in AII amacrine cells. Proteins that were three times or more abundant compared to the control condition were included in the string diagram.

Localization of BioID hits at Cx36 containing gap junctions in the AII amacrine cell

Earlier version: In the next set of experiments, we validated the most abundant candidates from our list in GFP expressing AII amacrine cells. We categorized the proteins based on the cell biological function they serve. In line with the mass spec data, we observed the most frequent colocalization for Cx36 with scaffold proteins including ZO-1, ZO-2, Cingulin and Sipa1l3 (Figure 3A, arrows). Although ZO-1, ZO-2 and Cingulin are known components of electrical synapses (34, 44), Sipa1l3 has not been shown before to associate with Cx36, yet it colocalized with the connexin to a similar extent as the ZO proteins, suggesting that it could be essential for synapse formation or stabilization. In addition to scaffolding proteins, we confirmed an association with several adapter proteins of the endocytic machinery, including Eps15l1, Snap91, Hip1r and Itsn1 (Figure 3B, arrows). In comparison to the scaffold proteins, however, the colocalization of Cx36 with each of these endocytic components was clearly less frequent and more heterogenous, which appears to reflect different stages in the life cycle of Cx36. As most proteins we identified in this initial screen broadly labeled the entire inner plexiform layer (IPL – the inner synaptic layer of the retina), we performed a colocalization analysis to exclude a random overlap with Cx36. We quantified the number of Cx36 plaques that contained Cgn, Sipa1l3, Eps15l1 and Hip1r (Figure 3C & D) and compared the extent of colocalization to a control in which we flipped the Cx36 channel horizontally. For each candidate, we found a significant reduction in colocalization when the original images were compared to the flipped control, indicating that the proteins we have identified are actual components of electrical synapses.

Revised version: In the next set of experiments, we validated the most abundant candidates from our list in GFP expressing AII amacrine cells. We categorized the proteins based on the cell biological function they serve. In line with the mass spec data, we observed the most frequent colocalization for Cx36 with scaffold proteins including ZO-1, ZO-2, cingulin and Sipa1l3 (Figure 3A, arrows). Although ZO-1, ZO-2 and cingulin are known components of electrical synapses (35, 45), Sipa1l3 has not been shown before to associate with Cx36, yet it colocalized with the connexin to a similar extent as the ZO proteins, suggesting that it could be essential for synapse formation or stabilization. In addition to scaffolding proteins, we confirmed an association with several adapter proteins of the endocytic machinery, including Eps15l1, Snap91, Hip1r and Itsn1 (Figure 3B, arrows). In comparison to the scaffold proteins, however, the colocalization of Cx36 with each of these endocytic components was clearly less frequent and more heterogenous, which appears to reflect different stages in the life cycle of Cx36. As most proteins we identified in this initial screen broadly labeled the entire inner plexiform layer (IPL – the inner synaptic layer of the retina), we performed a colocalization analysis to exclude a random overlap with Cx36. We quantified the number of Cx36 plaques that contained Cgn, Sipa1l3, Eps15l1 and Hip1r (Figure 3C & D) and compared the extent of colocalization to a control in which we flipped the Cx36 channel horizontally. For each candidate, we found a significant reduction in colocalization, when the original images were compared to the flipped control, indicating that the proteins we have identified are actual components of electrical synapses.

Localization of scaffold and endocytosis proteins identified by BioID at Cx36 gap junctions in AAV transduced AII amacrine cells expressing GFP.

(A) Colocalization of Cx36 and scaffold proteins including known interactors such as ZO-1, ZO-2 and cingulin. Besides these known interactors, we identified Sipa1l3, a PDZ domain containing protein implicated in cell adhesion and cytoskeletal organization. Sipa1l3 shows abundant colocalization with Cx36. (B) Colocalization of Cx36 and components of the endocytosis machinery. Among all proteins we have tested, we found frequent colocalization for EpsS15I1, an endocytic adapter protein, and Cx36. Scale: 10 µm. Magnified inset: 1 µm. (C) Bar graphs showing the degree of colocalization as the percentage of Cx36 puncta that overlap with Cgn and Sipa1l3. Colocalization was quantified for the original image and a horizontally flipped image to exclude random overlap. Bar graphs with SEM, p <0.05 *** for Cx36 and Cgn and p <0.0001 **** for Cx36 and Sipa1l3. Significance for Cx36 and Cgn was determined using a Wilcoxon matched-pairs signed rank test. Significance for Cx36 and Sipa1l3 was determined using a two-tailed paired t-test. (D) Bar graphs showing the degree of colocalization as the percentage of Cx36 puncta that overlap with endocytosis proteins the Ep15l1 and Hip1r. Bar graphs with SEM, p <0.0001 **** for Cx36 and Eps15l1 and p <0.05 *** for Cx36 and Hip1r. Significance was determined using a two-tailed paired t-test.

Earlier version: Besides synaptic scaffolds and endocytic adapter proteins we tested a variety of additional hits, which include synaptojanin 2 binding protein (Sj2bp), synaptotagmin 4 (categorized as “trafficking”) and several proteins that regulate actin dynamics such as G-protein regulated inducer of neurite outgrowth 1 (Gprin1), dedicator of cytokinesis 7 (Dock7) and microtubule associated protein 6 (Map6). All of these proteins showed considerable colocalization with Cx36 in AII amacrine cell dendrites (Figure 4A, B). We also tested components of chemical synapses including GluR2-3, the scaffold Shank2 (Figure 4D) and Bai1 (Figure 4C) a synaptic adhesion molecule implicated in excitatory synapse formation (49). We found that these proteins showed a partial overlap with Cx36 at the periphery of each junction.

Revised version: Besides synaptic scaffolds and endocytic adapter proteins we tested a variety of additional hits, which include synaptojanin 2 binding protein (Sj2bp), synaptotagmin 4 (Syt4; categorized as “trafficking”) and several proteins that regulate actin dynamics such as G-protein regulated inducer of neurite outgrowth 1 (Gprin1), dedicator of cytokinesis 7 (Dock7) and microtubule associated protein 6 (Map6). All of these proteins showed considerable colocalization with Cx36 in AII amacrine cell dendrites (Figure 4A, B). We also tested components of chemical synapses including GluR2-3, the scaffold Shank2 (Figure 4D) and Bai1 (Figure 4C) a synaptic adhesion molecule implicated in excitatory synapse formation (50). We found that these proteins showed a partial overlap with Cx36 at the periphery of each junction. In addition to the candidates described here, we identified additional candidates that were associated with Cx36 in AII amacrine cells (supplemental figure 1-4). These include Ajm1 (Apical junction component 1), Cap1 (Adenylyl cyclase associated protein1), Neurobeachin (Nbea), Ncam1 (Neural cell adhesion molecule 1), Rtn4 (Reticulon4), Af6 (Afadin), Trim9 (Tripartite motif containing protein 9), Golga4 (Golgin subfamily A member 4), Stxbp1 (Syntaxin binding protein 1) and Sipa1l1 (Signal-induced proliferation-associated 1-like protein 1) (Supplementary figures 1-5).

Localization of trafficking, cytoskeleton-associated, adhesion and synaptic proteins identified by BioID at Cx36 gap junctions in AAV transduced AII amacrine cells expressing GFP.

(A) Proteins implicated in membrane trafficking Sj2bp and Syt4 colocalize with Cx36 in AII amacrine cells. (B) Several cytoskeleton associated proteins and regulators such as Map6, Dock7 or Gprin1 colocalize with Cx36 in AII cell dendrites. (C) The adhesion molecule Bai1 was often colocalized with Cx36 in AII cell dendrites. (D) Often components of chemical synapses such as Shank2 and Glur2-3 were found in the periphery of gap junction plaques in AII amacrine cells. Scale: 10 µm. Magnified inset: 1 µm.

Added in the revised version: Sipa1l3, Sj2bp and Eps15I1 interact with Cx36.

Added in the revised version: (A) Coexpression of Sipa1l3 with ZO-1, ZO-2 and Cx36. In transfected HEK293T cells Sipa1l3 colocalizes with all three proteins. Colocalization of Sj2bp with Cx36 but not with a Cx36 mutant that lacks the PDZ binding motif. (B) Domain organization of constructs used for IP experiments. Co-precipitation of Cx36 and Sjbp2 and Sipa1l3 from lysates of co-transfected HEK293T cells. The PDZ binding deficient Cx36/S318Ter fails to bind Sjbp2 and Sipa1l3. IP experiments with ZO-1 and ZO-2 serve as a positive control. The proteins that were captured in each IP experiment are indicated with arrows in magenta. Cx36 dimers in IP samples are indicated with red arrows and monomers are indicated with blue arrows. (C) In transfected HEK293T cells, Cx36 is associated with endogenous proteins that were also detected in AII amacrine cells. Line scans next to each panel illustrate the close association of these proteins (57) with Cx36 (magenta). These proteins were also captured via GFP directed proximity biotinylation. Scale: 10 µm. Magnified inset: 1µm.

Identification of novel Cx36 binding partners

Earlier version: Thus far, we have gained significant insight into the Cx36 interactome and identified several new molecules as components of electrical synapses in AII amacrine cells. To understand how these molecules are organized at the synapse we tested protein-protein interactions in co-transfected HEK293T cells. Interestingly, two of the novel proteins we detected in AII cells contain PDZ domains and thus may directly bind to Cx36 like ZO-1 and ZO-2 (Figure 5B, GFP-Trap IP with both ZO proteins served as a positive control); these proteins are Sipa1l3 and Sj2bp. We found that Myc-Sipa1l3 interacted with all major components of electrical synapses: ZO-1, ZO-2 and Cx36 in co-transfected HEK293T cells. The overlap of Sipa1l3 and the ZO proteins was confined to cell cortices and to actin like fibers in the cytoplasm (Figure 5A). Additionally, we performed IP experiments and demonstrated that Cx36 binds to Sipa1l3 in a PDZ dependent manner. A truncated version of Cx36 lacking the PDZ binding motif (Cx36/S318Ter) failed to interact with Sipa1l3 in IP experiments. In co-transfected cells we observed that Sj2bp, which contains a single PDZ domain and a transmembrane domain (50), colocalized with Cx36 but not with the Cx36/S318Ter, indicating that the interaction of these proteins requires the PDZ domain. In IP experiments using FLAG tagged Sj2bp as a bait we were unable to coprecipitate Cx36, which was rather surprising as both proteins clearly colocalized in co-transfected HEK293T cells. One possible explanation for the lack of a detectable interaction is that Sj2bp displays a low affinity for the PDZ binding motif (PBM) of Cx36, as it is the case for ZO-1 (45). We have previously shown that binding of Cx36 to ZO-1 can be artificially enhanced by a simple deletion of amino acids 313-319 within the C-terminus, creating a new PBM with the following sequence: RTYV (51). We repeated the IP with this mutant and observed substantial binding to FLAG-Sj2bp (Figure 5B), suggesting that Cx36 and Sj2bp are compatible interactors. In addition to the candidates we overexpressed to verify direct binding interactions, we found that Cx36 colocalized with several endogenous HEK293T proteins that also occurred in our AII cell specific data set, including Eps15l1, Gprin1 and Sec22b (Figure 5C).

Revised version: Thus far, we have gained significant insight into the Cx36 interactome and identified several new molecules as components of electrical synapses in AII amacrine cells. To understand how these molecules are organized at the synapse we tested protein-protein interactions in co-transfected HEK293T cells. Interestingly, two of the novel proteins we detected in AII cells contain PDZ domains and thus may directly bind to Cx36 like ZO-1 and ZO-2 (Figure 5B, GFP-Trap IP with both ZO proteins served as a positive control); these proteins are Sipa1l3 and Sj2bp. We found that Myc-Sipa1l3 interacted with all major components of electrical synapses: ZO-1, ZO-2 and Cx36 in co-transfected HEK293T cells. The overlap of Sipa1l3 and the ZO proteins was confined to cell cortices and to actin like fibers in the cytoplasm (Figure 5A). Additionally, we performed IP experiments and demonstrated that Cx36 binds to Sipa1l3 in a PDZ dependent manner. A truncated version of Cx36 lacking the PDZ binding motif (Cx36/S318Ter) failed to interact with Sipa1l3 in IP experiments. In co-transfected cells we observed that Sj2bp, which contains a single PDZ domain and a transmembrane domain (51), colocalized with Cx36 but not with the Cx36/S318Ter, indicating that the interaction of these proteins requires the PDZ domain. In IP experiments using FLAG tagged Sj2bp as a bait we were unable to coprecipitate Cx36, which was rather surprising as both proteins clearly colocalized in co-transfected HEK293T cells. One possible explanation for the lack of a detectable interaction, is that Sj2bp displays a low affinity for the PDZ binding motif (PBM) of Cx36, as it is the case for ZO-1 (46). We have previously shown that binding of Cx36 to ZO-1 can be artificially enhanced by a simple deletion of amino acids 313-319 within the C-terminus creating a new PBM with the following sequence: RTYV (52). We repeated the IP with this mutant and observed substantial binding to FLAG-Sj2bp (Figure 5B), suggesting that Cx36 and SJ2BP are compatible interactors. In addition to the candidates we overexpressed to verify direct binding interactions, we found that Cx36 colocalized with several endogenous HEK293T proteins that also occurred in our AII cell specific data set, including Eps15l1, Gprin1 and Sec22b (Figure 5C).

Removed in the revised version: Sipa1l3, Sj2bp and Eps15I1 interact with Cx36.

Removed in the revised version: (A) Coexpression of Sipa1l3 with ZO-1, ZO-2 and Cx36. In transfected HEK293T cells Sipa1l3 colocalizes with all three proteins. Colocalization of Sj2bp with Cx36 but not with a Cx36 mutant that lacks the PDZ binding motif. (B) Domain organization of constructs used for IP experiments. Co-precipitation of Cx36 and Sjbp2 and Sipa1l3 from lysates of co-transfected HEK293T cells. The PDZ binding deficient Cx36/S318Ter fails to bind Sjbp2 and Sipa1l3. IP experiments with ZO-1 and ZO-2 serve as a positive control. The proteins that were captured in each IP experiment are indicated with arrows in magenta. Cx36 dimers in IP samples are indicated with red arrows and monomers are indicated with blue arrows. (C) In transfected HEK293T cells, Cx36 is associated with endogenous proteins that were also detected in AII amacrine cells. Line scans next to each panel illustrate the close association of these proteins (56) with Cx36 (magenta). These proteins were also captured via GFP directed proximity biotinylation. Scale: 10 µm. Magnified inset: 1µm.

Molecular architecture of the AII amacrine cell/On cone bipolar cell gap junction

Earlier version: AII amacrine cells form two different sets of gap junctions, with neighboring AII amacrine cells and with ON cone bipolar cells to convey neural signals that originate in rod photoceptors into the cone pathway (Illustrated in figure 6A) (9). Ultrastructural studies have revealed an asymmetry in the cytoplasm of AII/ON cone bipolar cell junctions (AII/ONCBC), which exhibits an electron dense structure (termed fluffy material) in the AII amacrine cell but not in the cone bipolar cell, hinting at differences in the molecular composition of both sites of the synapse (23, 52). Moreover, apart from accessory proteins that bind to the channel, AII/ONCBC gap junctions are likely to be distinct from AII/AII gap junctions since they recognize ON cone bipolar cells as the correct synaptic partner. These two characteristics of the AII/ONCBC gap junction can only be explained by a unique composition of the synapse proteome. We reasoned that the electron density in the AII cell could harbor several proteins we identified in this study. To test this hypothesis, we triple labelled AAV/HKamac_GFP infected retinas with antibodies to Scgn, Cx36 and the main candidates from our screen. With this strategy it was possible for us to visualize AII/ONCBC contacts (Figure 6 D, arrows) and the precise localization of each candidate. To illustrate the localization of AII/ONCBC gap junctions in a comprehensive way, we reconstructed the morphology of AII amacrine cells using serial section transmission EM imaging. As seen in figure 6B, most AII/ONCBC gap junctions (shown in yellow) were located above homologous AII/AII contacts (Cyan) and contained electron densities on both sides (Figure 6C).

Revised version: AII amacrine cells form two different sets of gap junctions, with neighboring AII amacrine cells and with ON cone bipolar cells to convey neural signals that originate in rod photoceptors into the cone pathway (Illustrated in figure 6A) (9). Ultrastructural studies have revealed an asymmetry in the cytoplasm of AII/ON cone bipolar cell junctions (AII/ONCBC), which exhibits an electron dense structure (termed fluffy material) in the AII amacrine cell but not in the cone bipolar cell, hinting at differences in the molecular composition of both sites of the synapse (23, 53). Moreover, apart from accessory proteins that bind to the channel, AII/ONCBC gap junctions are likely to be distinct from AII/AII gap junctions since they recognize ON cone bipolar cells as the correct synaptic partner. These two characteristics of the AII/ONCBC gap junction can only be explained by a unique composition of the synapse proteome. We reasoned that the electron density in the AII cell could harbor several proteins we identified in this study. To test this hypothesis, we triple labelled AAV/HKamac_GFP infected retinas with Scgn, Cx36 and the main candidates from our screen. With this strategy it was possible for us to visualize AII/ONCBC contacts (Figure 6 D, arrows) and the precise localization of each candidate. To illustrate the localization of AII/ONCBC gap junctions in a comprehensive way, we reconstructed the morphology of AII amacrine cells using serial section transmission EM imaging. As seen in figure 6B, most AII/ONCBC gap junctions (shown in yellow) were located above homologous AII/AII contacts (Cyan) and contained electron densities on both sides (Figure 6C).

Localization of BioID “hits” at AII amacrine cell/ON cone bipolar cell junctions.

Earlier version: (A) Cartoon illustrating the neurons involved in the primary rod pathway and the subcellular localization of AII/Cone bipolar cell junctions (dashed rectangle). (B) 3D reconstruction of AII amacrine cells from the RC2 connectomics dataset illustrating the localization of AII-AII gap junctions (cyan plaques in second image) and AII/ON Cone bipolar cell gap junctions (yellow plaques in second image). (C) Transmission electron micrographs of AII/ON Cone bipolar cell gap junctions. Pseudo colored: Gap junction (Cyan), AII cell (Magenta) and ON cone bipolar cell (Yellow). White box indicates area shown to right at 40,000X magnification (0.27 nm/px resolution) with 20 degree tilt. The density profile for the ROI indicated by the yellow box confirms the pentalaminar (dark-light-dark-light-dark) structure of a gap junction. Also note the asymmetric cytoplasmic densities on the AII versus ON Cone bipolar cell sides. Scale: 0.5 µm. (D) Several of the proteins we identified in AII amacrine cells colocalize with Cx36 at AII/cone bipolar cell junctions. The right plot for each panel depicts an intensity scan of a horizontal region of interest in the middle of each gap junction (arrow). Scale: 10 µm. Magnified inset: 1 µm. (E) High resolution scan highlighting the subsynaptic distribution of SIPA1L3. Scale: 5 µm. Magnified inset: 1 µm.

Revised version: (A) Cartoon illustrating the neurons involved in the primary rod pathway and the subcellular localization of AII/Cone bipolar cell junctions (dashed rectangle). (B) 3D reconstruction of AII amacrine cells from the RC2 connectomics dataset illustrating the localization of AII-AII gap junctions (cyan plaques in second image) and AII/ON Cone bipolar cell gap junctions (yellow plaques in second image). (C) EM micrographs of AII/ON Cone bipolar cell gap junctions. Pseudo colored: Gap junction (Cyan), AII cell (Magenta) and ON cone bipolar cell (Yellow). White box indicates area shown to right at 40000X magnification (0.27 nm/px resolution) with 20 degree tilt. The density profile for the ROI indicated by the yellow box confirms the pentalaminar (dark-light-dark-light-dark) structure of a gap junction. Also note the asymmetric cytoplasmic densities on the AII versus ON Cone bipolar cell sides. Scale: 0.5 µm. (D) Several of the proteins we identified in AII amacrine cells colocalize with Cx36 at AII/cone bipolar cell junctions. The right plot for each panel depicts an intensity scan of a horizontal region of interest in the middle of each gap junction (arrow). Scale: 10 µm. Magnified inset: 1 µm. (E) High resolution scan highlighting the subsynaptic distribution of SIPA1L3. Scale: 5 µm. Magnified inset: 1 µm.

Earlier version: The scaffold proteins ZO-1, ZO-2, Cgn and Sipa1l3 localized at AII/ONCBC contacts suggesting that they are a part of the electron density in AII cells (Figure 6D, arrows). In addition to these proteins, we also confirmed that AII/Cone bipolar cell gap junctions contained Syt4, Eps15l1, Sj2bp and Gprin1. Finally, we addressed the exact localization of the G protein coupled adhesion receptor Bai1, a synaptic adhesion molecule that has been implicated in excitatory synapse formation (49). Interestingly, we observed Bai1 at AII/ONCBC contacts, in the vicinity of Cx36, which is consistent with the notion of synaptic adhesion molecules as an integral part of electrical synapses (21, 53). This finding further implies that Bai1 might play a role in the formation of AII/ONCBC contacts. However, further knock-out studies will be necessary to address the concrete function of this adhesion molecule. Finally, we addressed the subcellular distribution of Sipa1l3 at AII/ONBC gap junctions. We noticed that these junctions often displayed a ring like shape, that contained Sipa1l3 in the center of each ring (Figure 6D). One possible interpretation of this distinct sub-synaptic distribution is that Sipa1l3 is interacting with a type of adhesion molecule that is directly adjacent to the gap junction.

Revised version: The scaffold proteins ZO-1, ZO-2, Cgn and Sipa1l3 localized at AII/ONCBC contacts suggesting that they are a part of the electron density in AII cells (Figure 6D, arrows). In addition to these proteins, we also confirmed that AII/Cone bipolar cell gap junctions contained Syt4, Eps15l1, Sj2bp and Gprin1. Finally, we addressed the exact localization of the G protein coupled adhesion receptor Bai1, a synaptic adhesion molecule that has been implicated in excitatory synapse formation (50). Interestingly, we observed Bai1 at AII/ONCBC contacts, in the vicinity of Cx36, which is consistent with the notion of synaptic adhesion molecules as an integral part of electrical synapses (21, 54). This finding further implies that Bai1 might play a role in the formation of AII/ONCBC contacts. However, further knock-out studies will be necessary to address the concrete function of this adhesion molecule. Finally, we addressed the subcellular distribution of Sipa1l3 at AII/ONBC gap junctions. We noticed that these junctions often displayed a ring like shape, that contained Sipa1l3 in the center of each ring (Figure 6E). One possible interpretation of this distinct sub-synaptic distribution is that Sipa1l3 is interacting with a type of adhesion molecule that is directly adjacent to the gap junction.

Electrical synapse scaffolds are targeted to AII amacrine cell/ ON cone bipolar cell contacts in the absence of Cx36

Earlier version: Previous studies have shown that ZO-1 localizes to neuronal contact sites even in the absence of connexins (33). This finding is consistent with a study by Meyer et al., (2014), who reported that Cx36-EGFP containing gap junctions colocalize with ZO-1 at AII/ONCBC contacts despite a C-terminal tag in the fusion protein that is known to interfere with PDZ domain mediated interactions (38, 51). One possible interpretation of these observations is that ZO-1, a protein with multiple protein-protein interaction domains, is connected to additional components of the gap junction such as adhesion molecules, the actin skeleton or other MAGUKs that retain the scaffold at the synapse. We wondered if a similar principle applies to the scaffold proteins that were identified in this study and tested the localization of ZO-2 and Sipa1l3 in Cx36 KO retinas that were infected with AAV/HKamac_GFP to visualize AII/ONCBC contacts. Here, we observed that ZO-2 and Sipa1l3 showed a similar effect as ZO-1 and still localized to AII/ONCBC contacts in Cx36 KO retinas (Figure 7A-C). Additionally, we quantified the density and size of immunoreactive Sipa1l3 and ZO-1 in the inner plexiform layer of WT and Cx36 KO mouse retina. While the size of ZO-2 and Sipa1l3 puncta were unchanged in Cx36 KO mice, the density of both ZO-2 and Sipa1l3 puncta were reduced and the colocalization of ZO-2 and Sipa1l3 was reduced in Cx36KO retina. Thus, loss of Cx36 resulted in a quantitative defect in formation of electrical synapse density complexes.

Revised version: Previous studies have shown that ZO-1 localizes to neuronal contact sites even in the absence of connexins (34). This finding is consistent with a study by Meyer et al., (2014), who reported that Cx36-EGFP containing gap junctions colocalize with ZO-1 at AII/ONCBC contacts despite a C-terminal tag in the fusion protein that is known to interfere with PDZ domain mediated interactions (39, 52). One possible interpretation of these observations is that ZO-1, a protein with multiple protein-protein interaction domains, is connected to additional components of the gap junction such as adhesion molecules, the actin skeleton or other MAGUKs that retain the scaffold at the synapse. We wondered if a similar principle applies to the scaffold proteins that were identified in this study and tested the localization of ZO-2 and Sipa1l3 in Cx36 KO retinas that were infected with AAV/HKamac_GFP to visualize AII/ONCBC contacts. Here, we observed that ZO-2 and Sipa1l3 showed a similar effect as ZO-1 and still localized to AII/ONCBC contacts in Cx36 KO retinas (Figure 7A-C). Additionally, we quantified the density and size of immunoreactive Sipa1l3 and ZO-1 in the inner plexiform layer of WT and Cx36 KO mouse retina (Figure 7D). While the size of ZO-2 and Sipa1l3 puncta were unchanged in Cx36 KO mice, the density of both ZO-2 and Sipa1l3 puncta were reduced and the colocalization of ZO-2 and Sipa1l3 was reduced in Cx36KO retina. Thus, loss of Cx36 resulted in a quantitative defect in formation of electrical synapse density complexes.

AII amacrine cell/ON cone bipolar cell contacts remain in Cx36 KO retinas and contain electrical synapse scaffolds.

Earlier version: AII amacrine cell/ON cone bipolar cell contacts were visualized with GFP and Scgn. In Cx36 KO retinas, ZO-1, ZO-2 and Sipa1l3 still localize to AII/CBC contacts. (D) Both Sipa1l3 and ZO-1 puncta density in the inner plexiform layer were reduced in Cx36KO, but puncta size was unchanged. Colocalization of Sipa1l3 and ZO-1 was reduced in Cx36KO. (E-F) Cartoon illustrating electrical synapse configurations of AII amacrine cells and identified components of the electrical synapse proteome of the AII amacrine cell. Scale: 10 µm. Magnified inset: 1µm.

Revised version: (A-C) AII amacrine cell/ON cone bipolar cell contacts were visualized with GFP and SCGN. In Cx36 KO retinas, ZO-1, ZO-2 and Sipa1l3 still localize to AII/CBC contacts. (D) Both Sipa1l3 and ZO-1 puncta density in the inner plexiform layer were reduced in Cx36KO, but puncta size was unchanged. Colocalization of Sipa1l3 and ZO-1 was reduced in Cx36KO. (E-F) Cartoon illustrating Scale: 10 µm. Magnified inset: 1µm.

Discussion

Earlier version: In the present study, we have used TurboID to uncover the electrical synapse proteome in retinal neurons. Our screen identified a plethora of molecules that were associated with the neuronal connexins Cx36/Cx35b in zebrafish and mice such as adhesion molecules, scaffold proteins, chemical synapse proteins, components of the endocytic machinery and cytoskeleton associated proteins (Illustrated in Figure 7 E-F). The presence of ZO proteins and endocytosis components in the Cx36 and the Cx35b interactome data suggests a certain degree of conservation in the proteome of electrical synapses. An apparent difference between species, however, is the existence of additional ZO variants and Cx36 homologues in the zebrafish genome (33), which are also represented in our Cx35b interactome. In contrast to the bass retina, where it was shown that Cx35 and Cx34.7 form two distinct circuits (35), we found that zebrafish photoreceptors formed gap junctions that contained both connexins. Cell culture experiments have confirmed that Cx35 and Cx34.7 are able to form heterotypic channels (54) that exhibit slightly altered voltage sensitivities compared to the homotypic channels. How exactly these gap junction channels in zebrafish photoreceptors are configured and which exact cell types they connect remains to be determined.

Revised version: In the present study, we have used TurboID to uncover the electrical synapse proteome in retinal neurons. Our screen identified a plethora of molecules that were associated with the neuronal connexins Cx36/Cx35b in zebrafish and mice such as adhesion molecules, scaffold proteins, chemical synapse proteins, components of the endocytic machinery and cytoskeleton associated proteins (Illustrated in Figure 7 E-F). The presence of ZO proteins and endocytosis components in the Cx36 and the Cx35b interactome data suggests a certain degree of conservation in the proteome of electrical synapses (Supplementary Table 3). An apparent difference between species, however, is the existence of additional ZO variants and Cx36 homologues in the zebrafish genome (34), which are also represented in our Cx35b interactome. In contrast to the bass retina, where it was shown that Cx35 and Cx34.7 form two distinct circuits (36), we found that zebrafish photoreceptors formed gap junctions that contained both connexins. Cell culture experiments have confirmed that Cx35 and Cx34.7 are able to form heterotypic channels (55) that exhibit slightly altered voltage sensitivities compared to the homotypic channels. How exactly these gap junction channels in zebrafish photoreceptors are configured and which exact cell types they connect remains to be determined.

Earlier version: Among other proteins that were detected in this study, we identified Sipa1l3 as a novel scaffold protein for electrical synapses. Interestingly, Sipa1l3 was also described as a component of post synaptic densities of glutamatergic synapses in hippocampal neurons, and is known interact with Fezzins, which occur in a complex with Shank3 (55). Additionally, we found that Sipa1l3 was able to interact with Cx36 and the ZO proteins. This raises the question of what potential function Sipa1l3 might serve given the abundance of colocalization with Cx36 in AII amacrine cells? Genetic screens have shown that mutations in the Sipa1l3 gene result in abnormal eye and lens development and a decrease in the formation of cell adhesions. Like other members of the Sipa1l family, Sipa1l3 also contains a RAP GTPase activating protein (GAP) domain that is known to regulate the activity of RAP proteins, which in the GTP-bound form impacts processes such as adhesion and actin dynamics (56). One way in which Sipa1l3 might influence the formation of electrical synapses might be as a regulator of Rap1 and its effector proteins. To determine the exact function of Sipa1l3 further studies on KO mice will be required.

Revised version: Among other proteins that were detected in this study, we identified Sipa1l3 as a novel scaffold protein for electrical synapses. Interestingly, Sipa1l3 was also described as a component of post synaptic densities of glutamatergic synapses in hippocampal neurons, and is known interact with Fezzins, which occur in a complex with Shank3 (56). Additionally, we found that Sipa1l3 was able to interact with Cx36 and the ZO proteins. This raises the question of what potential function Sipa1l3 might serve given the abundance of colocalization with Cx36 in AII amacrine cells? Genetic screens have shown that mutations in the Sipa1l3 gene result in abnormal eye and lens development and a decrease in the formation of cell adhesions. Like other members of the Sipa1l family, Sipa1l3 also contains a RAP GTPase activating protein (GAP) domain that is known to regulate the activity of RAP proteins, which in the GTP-bound form impacts processes such as adhesion and actin dynamics (57). One way in which Sipa1l3 might influence the formation of electrical synapses might be as a regulator of Rap1 and its effector proteins. To determine the exact function of Sipa1l3 further studies on KO mice will be required.

Earlier version: One striking aspect of the Cx36 interactome in AII amacrine cells is the abundance of the endocytosis machinery. Considering the short half life time of Cx36 of 3.1 h (29) and the fact that the turnover of electrical synapse proteins is a steady-state process (43), it makes sense that the endocytic machinery is that well represented in our interactome data. Thus, besides phosphorylation, it seems possible that the cellular regulation of turnover mechanisms could function as an additional means to adjust the strength of electrical coupling between AII amacrine cells.

Revised version: One striking aspect of the Cx36 interactome in AII amacrine cells is the abundance of the endocytosis machinery. Considering the short half life time of Cx36 of 3.1 h (29) and the fact that the turnover of electrical synapse proteins is a steady-state process (44), it makes sense that the endocytic machinery is that well represented in our interactome data. Thus, besides phosphorylation, it seems possible that the cellular regulation of turnover mechanisms could function as an additional means to adjust the strength of electrical coupling between AII amacrine cells.

Is Bai1 necessary to form electrical synapses in the primary rod pathway?

Earlier version: How and when it is determined where electrical synapses are formed is currently unknown. It has been suggested that neuronal adhesion molecules serve as cues that wire the correct neurons together (3). This task also has to be accomplished in the primary rod pathway, where AII amacrine cells form synapses among each other and with ON Cone bipolar cells. Interestingly, we identified the brain specific angiogenesis inhibitor 1 (Bai1) between AII amacrine and ON cone bipolar cell gap junctions, which could suggest that Bai1 is necessary to connect these cells. Recent reports have described Bai1 as a synaptogenic enzyme capable of regulating intracellular signaling pathways that determine synapse formation. (49, 57). Thus, besides merely functioning as a cell adhesion molecule, Bai1 might actively support synaptogenesis in AII amacrine cells at a level that is hierarchical to Cx36 and its associated scaffolds, which would explain why these cells form contacts with ON Cone bipolar cells even in the absence of Cx36.

Revised version: How and when it is determined where electrical synapses are formed is currently unknown. It has been suggested that neuronal adhesion molecules serve as cues that wire the correct neurons together (3). This task also has to be accomplished in the primary rod pathway, where AII amacrine cells form synapses among each other and with ON Cone bipolar cells. Interestingly, we identified the brain specific angiogenesis inhibitor 1 (Bai1) between AII amacrine and ON cone bipolar cell gap junctions, which could suggest that Bai1 is necessary to connect these cells. Recent reports have described Bai1 as a synaptogenic enzyme capable of regulating intracellular signaling pathways that determine synapse formation. (50, 58). Thus, besides merely functioning as a cell adhesion molecule, Bai1 might actively support synaptogenesis in AII amacrine cells at a level that is hierarchical to Cx36 and its associated scaffolds, which would explain why these cells form contacts with ON Cone bipolar cells even in the absence of Cx36.

Cell type-specific differences in the Cx36 interactome?

Earlier version: Interestingly, our interactome in AII amacrine cells showed no signs of the scaffolds Mupp1 and Af6, although these candidates were shown to colocalize and interact with Cx36 in mammalian retinas and several other brain regions (31). Moreover, both proteins also showed up in our Cx35b interactome in zebrafish retina. One feasible explanation is that AII cell gap junctions simply lack these scaffolds and that they may instead be used for electrical synapses of other neurons like bipolar cells, which our approach did not consider. In general, it will be interesting to understand how well certain components of the electrical synapse proteome are preserved between cell types. Endocytosis and ZO proteins, for instance, are unlikely to be cell-type specific interactors of gap junction channels and are more likely to be ubiquitous constituents of electrical synapses, given their critical function in the constant turnover of Cx36. This is also reflected in the evolutionary conservation of these components in the proteomes from zebrafish and mice (supplementary figure 3). Bai1, on the other hand, a protein that is needed to establish specific synaptic contacts, may only occur in certain electrical synapses.

Revised version: Our Cx36 interactome in AII amacrine cells showed no signs of the scaffold protein Mupp1, although it was previously identified as component of electrical synapses in the retina and several other brain regions (32). Interestingly, however, Mupp 1 was detected in our Cx35b interactome in zebrafish retina. One feasible explanation is that AII cell gap junctions simply lack certain scaffolds such as Mupp 1 that are used by electrical synapses of other neurons like bipolar cells, which our approach did not consider. In general, it will be interesting to understand how well certain components of the electrical synapse proteome are preserved between cell types. Endocytosis and ZO proteins, for instance, are unlikely to be cell-type specific interactors of gap junction channels and are more likely to be ubiquitous constituents of electrical synapses, given their critical function in the constant turnover of Cx36. This is also reflected in the evolutionary conservation of these components in the proteomes from zebrafish and mice. Bai1, on the other hand, a protein that is needed to establish specific synaptic contacts, may only occur in certain electrical synapses.

Material and methods

Animal husbandry of zebrafish

Earlier version: Maintenance and breeding of zebrafish was conducted under standard conditions (58). Wild type (AB) zebrafish were purchased from the Zebrafish International Resource Center (ZIRC, Eugene, OR, USA). Fish were maintained on a 14 h light/10 h dark cycle. All procedures were performed in accordance with the ARVO statement on the use of animals in ophthalmic and vision research and US Public Health Service guidelines and have been reviewed and approved by the Institutional Animal Care and Use Committees at the University of Texas Health Science Center at Houston and the University of Houston.

Revised version: Maintenance and breeding of zebrafish was conducted under standard conditions (59). Wild type (AB) zebrafish were purchased from the Zebrafish International Resource Center (ZIRC, Eugene, OR, USA). Fish were maintained on a 14 h light/10 h dark cycle. All procedures were performed in accordance with the ARVO statement on the use of animals in ophthalmic and vision research and US Public Health Service guidelines and have been reviewed and approved by the Institutional Animal Care and Use Committees at the University of Texas Health Science Center at Houston and the University of Houston.

Tol2 mediated transgenesis

Earlier version: Transgenic Cx35b-V5-TurboID fish were generated using the Tol2 system and a donor vector containing 4.5 kb of the Cx35b promoter, exon1 of the Cx35b gene and the open reading frame of Cx35-V5-TurboID. The coding sequence of V5-TurboID (25) was cloned into a region of the Cx35b gene encoding the C-terminal tail. V5-TurboID was inserted between leucine 285 and proline 286 to expose the 19 C-terminal amino acids of Cx35b. Additionally, the donor vector contained a Myl7 promoter driving mCherry expression in the heart to select larvae in which the construct was integrated into genome. Transgenic Cx34.7-oxGFP fish were generated in the same way using a donor vector containing 2.8 kb of the Cx34.7 promoter and the coding sequence of Cx34.7 with oxGFP inserted between aspartate 283 and methionine 284 in the C-terminus. To generate transgenic fish, the donor vector was co-injected with Tol2 transposase mRNA into embryos in the 1 cell stage as described previously (59). Transgenic individuals were raised to adulthood and mated.

Revised version: Transgenic Cx35b-V5-TurboID fish were generated using the Tol2 system and a donor vector containing 4.5 kb of the Cx35b promoter, exon1 of the Cx35b gene and the open reading frame of Cx35-V5-TurboID. The coding sequence of V5-TurboID (25) was cloned into a region of the Cx35b gene encoding the C-terminal tail. V5-TurboID was inserted between leucine 285 and proline 286 to expose the 19 C-terminal amino acids of Cx35b. Additionally, the donor vector contained a Myl7 promoter driving mCherry expression in the heart to select larvae in which the construct was integrated into genome. Transgenic Cx34.7-oxGFP fish were generated in the same way using a donor vector containing 2.8 kb of the Cx34.7 promoter and the coding sequence of Cx34.7 with oxGFP inserted between aspartate 283 and methionine 284 in the C-terminus. To generate transgenic fish, the donor vector was co-injected with Tol2 transposase mRNA into embryos in the 1 cell stage as described previously (60). Transgenic individuals were raised to adulthood and mated.

Animal Husbandry of mice

Earlier version: The experiments in this study were conducted with wildtype mice (C57BL/6J), the Cx36-EGFP strain (37–39) (Kindly provided by Dr. Hannah Monyer) and Cx36 Knock out mice (60) (Kindly provided by Dr. David Paul). Animals of ages 2-12 months were used for experiments. All procedures were performed in accordance with the ARVO statement on the use of animals in ophthalmic and vision research and US Public Health Service guidelines and were approved by the Institutional Animal Care and Use Committee at the University of Houston.

Revised version: The experiments in this study were conducted with wildtype mice (C57BL/6J), the Cx36-EGFP strain (38–40) (Kindly provided by Dr. Hannah Monyer) and Cx36 Knock out mice (61) (Kindly provided by Dr. David Paul). Animals of ages 2-12 months were used for experiments. All procedures were performed in accordance with the ARVO statement on the use of animals in ophthalmic and vision research and US Public Health Service guidelines and were approved by the Institutional Animal Care and Use Committee at the University of Houston.

AAV constructs and intravitreal injections

Earlier version: All AAV constructs used in this study were generated with the pAAV-9-(5)-hSYN-CAMKII-GFP vector kindly provided by Dr. Neal Waxham. The coding sequence of the AII amacrine cell specific AII promoter Hkamac was previously described by Khabou et al., (2023), and synthesized as a 654 base pair gBlocks® fragment by IDT genomics. The hSYN promoter and the coding sequence of CaMKII in the pAAV-9-(5)-hSYN-CAMKII-GFP vector were excised via XbaI and EcoRV and the coding sequence of V5-TurboID-dGBP (36, 51, 61) and GFP were fused to the AII promoter and integrated into the vector via Gibson assembly. Titers of AAV particles for different AAV constructs varied ranging from 8×1012 to 1×1013 gc/ml. For each experimental condition similar amounts of virus particles were injected. All virus particles were generated at the Baylor College of Medicine gene vector core facility. All AAV vectors that were used in this study will be made available via Addgene. Prior to each virus injection, mice were anesthetized with 3.5-3.9% isoflurane using the SomnoSuite/Mouse STAT Pulse Oximeter (SS-01-03, Kent Scientific Corporation). Once the animal was anesthetized, a small incision next to the pupil was made using a 30 G needle. Afterwards, 1.5 µl of the virus solution was carefully injected intravitreally. Eyes were collected 2-4 weeks after each experiment.

Revised version: All AAV constructs used in this study were generated with the pAAV-9-(5)-hSYN-CAMKII-GFP vector kindly provided by Dr. Neal Waxham. The coding sequence of the AII amacrine cell specific AII promoter Hkamac was previously described by Khabou et al., (2023), and synthesized as a 654 base pair gBlocks® fragment by IDT genomics. The hSYN promoter and the coding sequence of CaMKII in the pAAV-9-(5)-hSYN-CAMKII-GFP vector were excised via XbaI and EcoRV and the coding sequence of V5-TurboID-dGBP (37, 52, 62) and GFP were fused to the AII promoter and integrated into the vector via Gibson assembly. Titers of AAV particles for different AAV constructs varied ranging from 8×1012 to 1×1013 gc/ml. For each experimental condition similar amounts of virus particles were injected. All virus particles were generated at the Baylor College of Medicine gene vector core facility. All AAV vectors that were used in this study will be made available via Addgene. Prior to each virus injection, mice were anesthetized with 3.5-3.9% isoflurane using the SomnoSuite/Mouse STAT Pulse Oximeter (SS-01-03, Kent Scientific Corporation). Once the animal was anesthetized, a small incision next to the pupil was made using a 30 G needle. Afterwards, 1.5 µl of the virus solution was carefully injected intravitreally. Eyes were collected 2-4 weeks after each experiment.

DNA constructs

Earlier version: FLAG-SJ2BP and Myc-SIPA1L3 constructs were used in previous studies (50, 62). Venus-ZO-1 (#56394) and GFP-ZO-2 (#27422) were purchased from Addgene. The FLAG-Gprin1 (Clone: OMu09339D) and FLAG-EPS15L1 (Clone: OMu22488D) constructs were purchased from Genscript. All Cx36 pcDNA3.1 expression vectors were generated in previous studies (51, 61, 63).

Revised version: FLAG-SJ2BP and Myc-SIPA1L3 constructs were used in previous studies (51, 63). Venus-ZO-1 (#56394) and GFP-ZO-2 (#27422) were purchased from Addgene. The FLAG-Gprin1 (Clone: OMu09339D) and FLAG-EPS15L1 (Clone: OMu22488D) constructs were purchased from Genscript. All Cx36 pcDNA3.1 expression vectors were generated in previous studies (52, 62, 64).

In vivo Biotinylation

Removed in the revised version: In vivo biotinylation in zebrafish

Earlier version: Prior to biotin injections, adult zebrafish were anesthetized with 0.02% Tricaine and immobilized in a small wax mold. A small incision was made into the abdomen and 30 µl of a 5mM biotin (solved in PBS) solution was injected intraperitoneally using a Hamilton syringe. This procedure was repeated for 2 consecutive days. Afterwards the fish were sacrificed by a cold shock and the retinas were extracted. The streptavidin pull down was carried out as previously described (51, 61). 80 retinas for each condition (Wild type vs. Cx35b-V5-TurboID) were used.

Revised version: In vivo biotinylation in zebrafish. Prior to biotin injections, adult zebrafish were anesthetized with 0.02% Tricaine and immobilized in a small wax mold. A small incision was made into the abdomen and 30 µl of a 5mM biotin (solved in PBS) solution was injected intraperitoneally using a Hamilton syringe. This procedure was repeated for 2 consecutive days. Afterwards the fish were sacrificed by a cold shock and the retinas were extracted. The streptavidin pull down was carried out as previously described (52, 62). 80 retinas for each condition (Wild type vs. Cx35b-V5-TurboID) were used. The experiment was performed twice.

Removed in the revised version: In vivo biotinylation in mouse

Earlier version: 3 ½ weeks after intravitreal virus injections mice were anesthetized with isoflurane and 1ml of a 5mM biotin was injected subcutaneously for 4 consecutive days. 17 retinas of each condition (Cx36-EGFP vs. wild type) were used for the streptavidin pull down.

Revised version: In vivo biotinylation in mouse. 3 ½ weeks after intravitreal virus injections mice were anesthetized with isoflurane and 1ml of a 5mM biotin was injected subcutaneously for 4 consecutive days. 17 retinas of each condition (Cx36-EGFP vs. wild type) were used for the streptavidin pull down. The experiment was performed twice.

LC/MS/MS analysis

For mass spectrometry analysis of streptavidin pull downs, samples were concentrated on a acrylamide/bis-acrylamide stacking gel and processed via In-gel digestion. An aliquot of the tryptic digest (in 2% acetonitrile/0.1% formic acid in water) was analyzed by LC/MS/MS on an Orbitrap FusionTM TribridTM mass spectrometer (Thermo ScientificTM) interfaced with a Dionex UltiMate 3000 Binary RSLCnano System. Peptides were separated onto an analytical C18 column (100μm ID x 25 cm, 5 μm, 18Å Reprosil-Pur C18-AQ beads from Dr Maisch, Ammerbuch-Entringen, Germany) at flow rate of 350 nl/min. Gradient conditions were: 3%-22% B for 90 min; 22%-35% B for 10min; 35%-90% B for 10 min; 90% B held for 10 min, (solvent A, 0.1 % formic acid in water; solvent B, 0.1% formic acid in acetonitrile). The peptides were analyzed using data-dependent acquisition method, Orbitrap Fusion was operated with measurement of FTMS1 at resolutions 120,000 FWHM, scan range 350-1500 m/z, AGC target 2E5, and maximum injection time of 50 ms; During a maximum 3 second cycle time, the ITMS2 spectra were collected at rapid scan rate mode, with HCD NCE 34%, 1.6 m/z isolation window, AGC target 1E4, maximum injection time of 35 ms. Dynamic exclusion was employed for 20 seconds.

Data processing and analysis

Earlier version: The raw data files were processed using Thermo Scientific™ Proteome Discoverer™ software version 1.4. Spectra were searched against the Uniprot-Mouse or Uniprot-Zebrafish database using Sequest. Search results were filtered to 1% FDR for strict conditions and 5% for relaxed conditions using Percolator. Trypsin was set as the enzyme with maximum missed cleavages set to 2. For peptide mapping, MS tolerance was set to 10 ppm, and MS/MS tolerance to 0.6 Da. Carbamidomethylation of cysteine residues was used as a fixed modification, while oxidation of methionine, N-terminal acetylation, and phosphorylation of serine and threonine were set as variable modifications.

Revised version: The raw data files were processed using Thermo Scientific™ Proteome Discoverer™ software version 1.4. Spectra were searched against the Uniprot-Zebrafish or Uniprot-Mouse database using Sequest. Search results were filtered to 1% FDR for strict conditions and 5% for relaxed conditions using Percolator. Trypsin was set as the enzyme with maximum missed cleavages set to 2. For peptide mapping. MS tolerance was set to 10 ppm, and MS/MS tolerance to 0.6 Da. Carbamidomethylation of cysteine residues was used as a fixed modification, while oxidation of methionine, N-terminal acetylation, and phosphorylation of serine and threonine were set as variable modifications.

Bioinformatics

Earlier version: A ratio of Cx36 conditions to the control was used to identify any proteins that were enriched greater than 3-fold in mouse retina and 2-fold in zebrafish retina. Proteins above this threshold were run through the Cytoscape (V3.10.2) plugin ClueGo (V2.5.10) to identify the GO Biological Processes that were upregulated in each Cx36 condition across both zebrafish and mouse species (64, 65). Pathways that show a p-value of less than 0.05 were used for analysis. To determine protein-protein interaction networks, proteins above threshold were also run through the standard STRING (V12.0) pipeline (66).

Revised version: For both zebrafish and mouse, data from two experiments with matched experimental and control samples were integrated into a single master table. A mean ratio of Cx35b/Cx36 conditions to the control was used to identify proteins that were enriched greater than 3-fold. The data were also analyzed using Significance Analysis of Interactome (SAINT) (31) implemented via the Resource for Evaluation of Protein Interaction Networks (REPRINT), accessed at https://reprint-apms.org/. Proteins exceeding the 3-fold enrichment threshold and a SAINT score greater than 0.5 were run through the Cytoscape (V3.10.2) plugin ClueGo (V2.5.10) to identify the GO Biological Processes that were upregulated in each Cx36 condition across both zebrafish and mouse species (65, 66). Pathways that show a p-value of less than 0.05 were used for analysis. To determine protein-protein interaction networks, proteins above threshold were also run through the standard STRING (V12.0) pipeline (67).

Cell culture

Earlier version: Human embryonic kidney 293 T cells (HEK293T/17; catalog #CRL-11268; ATCC, Manassas, VA, USA) were grown in Dulbecco’s Modified Eagle Medium (DMEM) supplemented with 10% fetal bovine serum (FBS), 1% penicillin and streptomycin, and 1% non-essential amino acids (all Thermo Fisher Scientific, Rockford, IL, USA) at 37 °C in a humidified atmosphere with 5% CO2. For pull-down experiments 1 million cells were plated in 60mm dishes. Transfections were carried out with Geneporter 2 as previously described (51).

Revised version: Human embryonic kidney 293 T cells (HEK293T/17; catalog #CRL-11268; ATCC, Manassas, VA, USA) were grown in Dulbecco’s Modified Eagle Medium (DMEM) supplemented with 10% fetal bovine serum (FBS), 1% penicillin and streptomycin, and 1% non-essential amino acids (all Thermo Fisher Scientific, Rockford, IL, USA) at 37 °C in a humidified atmosphere with 5% CO2. For pull-down experiments 1 million cells were plated in 60mm dishes. Transfections were carried out with Geneporter 2 as previously described (52).

Immunocytochemistry

Earlier version: Immunolabeling of transfected HEK293T cells was carried out as previously described (61, 63). Transfected HEK293T expressing the proteins of interest were briefly rinsed in PBS and fixed in 2% paraformaldehyde (PFA) solved in PBS for 15 min at RT. After the fixation step, the coverslips were washed three times in PBS for 10 min and incubated with the primary antibody solution containing 10% normal donkey serum in 0.5% Triton X-100 in PBS overnight at 4 °C .The following antibodies were used: mouse anti Cx36, 1:500 (MAB3045, Millipore), rabbit anti SIPA1l3, 1:200 (30544-1-AP, Proteintech), anti GFP, 1:250 (A10262, Thermofisher), rabbit anti EPS15L1, 1:250 (PA5-65940, Thermofisher) and mouse anti V5, 1:500 (R960-2, Thermofisher), rabbit anti myc (16286-1-AP, Proteintech). At the next day, the coverslips were washed 3×10 min in PBS and incubated with the secondary antibodies diluted in 10% normal donkey serum in 0.5% Triton X-100 in PBS at RT under light protected conditions. The following secondary antibodies were used: from donkey, 1:500, conjugated with Cy3, Alexa488, Alexa568, or Alexa647 (Jackson Immunoresearch, West Grove, PA). Afterwards coverslips were washed 3×10 min in PBS and mounted with Vectashield containing DAPI (H-2000, Vector Laboratories Inc.)

Revised version: Immunolabeling of transfected HEK293T cells was carried out as previously described (62, 64). Transfected HEK293T expressing the proteins of interest were briefly rinsed in PBS and fixed in 2% paraformaldehyde (PFA) solved in PBS for 15 min at RT. After the fixation step, the coverslips were washed three times in PBS for 10 min and incubated with the primary antibody solution containing 10% normal donkey serum in 0.5% Triton X-100 in PBS overnight at 4 °C .The following antibodies were used: mouse anti Cx36, 1:500 (MAB3045, Millipore), rabbit anti SIPA1l3, 1:200 (30544-1-AP, Proteintech), anti GFP, 1:250 (A10262, Thermofisher), rabbit anti EPS15L1, 1:250 (PA5-65940, Thermofisher) and mouse anti V5, 1:500 (R960-2, Thermofisher), rabbit anti Myc (16286-1-AP, Proteintech). At the next day, the coverslips were washed 3×10 min in PBS and incubated with the secondary antibodies diluted in 10% normal donkey serum in 0.5% Triton X-100 in PBS at RT under light protected conditions. The following secondary antibodies were used: from donkey, 1:500, conjugated with Cy3, Alexa488, Alexa568, or Alexa647 (Jackson Immunoresearch, West Grove, PA). Afterwards coverslips were washed 3×10 min in PBS and mounted with Vectashield containing DAPI (H-2000, Vector Laboratories Inc.)

Immunohistochemistry

Earlier version: Immunolabeling of vertical sections was carried out as previously described (67, 68). Mice were anesthetized with isoflurane and sacrificed via cervical dislocation. The eyes were removed from the animal and opened with a circular cut around the ora serrata. The lens was removed, eyes were fixed with 2% PFA solved in PBS for 20 min at room temperature and washed 3×10 min in PBS. Afterwards eyecups were stored in 30% sucrose at 4°C overnight. On the next day eyecup preparations were embedded in Tissue-Tek O.C.T. cryomatrix (Sakura Finetek, Torrance, CA, USA) and stored at -20°C. Eyecups were cut into 20 µm thin section and incubated at 37°C for 30 min. Afterwards the slides were washed with PBS (2×5 min) and incubated in the primary antibody solution containing 10% normal donkey serum in 0.5% Triton X-100 in PBS overnight. The following primary antibodies were used: Cx36, 1:250 (Clone: 8F6, MAB3045, Millipore); Cx36, 1:250 (Clone: 1E5H5, 37-4600, Thermofisher); GFP, 1:250 (A10262, Thermofisher); ZO-1, 1:250 (Clone: ZO1-1A12, 33-9100, Thermofisher); ZO-2, 1:250 (71-1400, Thermofisher); CGN, 1:250 (PA5-5561, Thermofisher); SIPA1L3, 1:100 (30544-1-AP, Proteintech); EPS15L1, 1:250 (PA5-65940, Thermofisher); HIP1R, 1:250 (AB9882,Sigma); GluR2-3, 1:200 (07-598, Sigma); Gprin1, 1:100 (13771-1-AP, Proteintech); DOCK7, 1:100 (13000-1-AP, Proteintech); MAP6, 1:100 (NBP2-14220, Novus Biologicals); Synaptotagmin4, 1:100 (105 143, Synaptic Systems); SJ2BP, 1:100 (15666-1-AP, Proteintech); BAI1, 1:100 (NB110-81586, Novus Biologicals); SHANK2, 1:500 (Synaptic Systems, 162204); SCGN, 1:250 (CSB-PA020821LA01HU, ARP); SCGN, 1:500 (RD184120100, Biovendor); Vglut1, 1:500 (AB5905, Sigma); ITNS1, 1:100-250 (PA5-115432, Thermofisher). The next day, the slides were washed 3×10 min and incubated with secondary antibodies diluted in 10% normal donkey serum in 0.5% Triton X-100 in PBS for 1h at RT under light protected conditions. The following secondary antibodies were used: from donkey, 1:250, conjugated with Cy3, Alexa488, Alexa568, or Alexa647 or Dylight 405 (Jackson Immunoresearch, West Grove, PA). Afterwards the slides were washed 3×10 min and mounted with Vectashield (H-1900, Vector Laboratories, Inc.).

Revised version: Immunolabeling of vertical sections was carried out as previously described (68, 69). Mice were anesthetized with isoflurane and sacrificed via cervical dislocation. The eyes were removed from the animal and opened with a circular cut around the ora serrata. The lens was removed, eyes were fixed with 2% PFA solved in PBS for 20 min at room temperature and washed 3×10 min in PBS. Afterwards eyecups were stored in 30% sucrose at 4°C overnight. On the next day eyecup preparations were embedded in Tissue-Tek O.C.T. cryomatrix (Sakura Finetek, Torrance, CA, USA) and stored at -20°C. Eyecups were cut into 20 µm thin section and incubated at 37°C for 30 min. Afterwards the slides were washed with PBS (2×5 min) and incubated in the primary antibody solution containing 10% normal donkey serum in 0.5% Triton X-100 in PBS overnight. The following primary antibodies were used: Cx36, 1:250 (Clone: 8F6, MAB3045, Millipore); Cx36, 1:250 (Clone: 1E5H5, 37-4600, Thermofisher); GFP, 1:250 (A10262, Thermofisher); ZO-1, 1:250 (Clone: ZO1-1A12, 33-9100, Thermofisher); ZO-2, 1:250 (71-1400, Thermofisher); CGN, 1:250 (PA5-5561, Thermofisher); SIPA1L3, 1:100 (30544-1-AP, Proteintech); EPS15L1, 1:250 (PA5-65940, Thermofisher); HIPR1, 1:250 (AB9882,Sigma); GluR2-3, 1:200 (07-598, Sigma); Gprin1, 1:100 (13771-1-AP, Proteintech); DOCK7, 1:100 (13000-1-AP, Proteintech); MAP6, 1:100 (NBP2-14220, Novus Biologicals); Synaptotagmin4, 1:100 (105 143, Synaptic Systems); SJ2BP, 1:100 (15666-1-AP, Proteintech); BAI1, 1:100 (NB110-81586, Novus Biologicals); SHANK2, 1:500 (Synaptic Systems, 162204); SCGN, 1:250 (CSB-PA020821LA01HU, ARP); SCGN, 1:500 (RD184120100, Biovendor); Vglut1, 1:500 (AB5905, Sigma); ITNS1, 1:100-250 (PA5-115432, Thermofisher). The next day, the slides were washed 3×10 min and incubated with secondary antibodies diluted in 10% normal donkey serum in 0.5% Triton X-100 in PBS for 1h at RT under light protected conditions. The following secondary antibodies were used: from donkey, 1:250, conjugated with Cy3, Alexa488, Alexa568, or Alexa647 or Dylight 405 (Jackson Immunoresearch, West Grove, PA). Afterwards the slides were washed 3×10 min and mounted with Vectashield (H-1900, Vector Laboratories, Inc.).

Immunoprecipitation

Earlier version: Immunoprecipitation experiments were carried out as previously described (46). 1.5 x106 HEK293T cells were plated on 60mm dishes and co-transfected at the next day with 2 µg of each vector using Geneporter2 (Amsbio) transfection reagent. 24h after transfection, cells were rinsed in phosphate buffered saline (pH 7.4) and transferred to a reaction tube. The cells were centrifuged at 5000g for 5 min and afterwards homogenized in immunoprecipitation buffer (IP buffer) containing 200mM NaCl, 2m EDTA, 1% Tx-100, 50mM Tris HCL (pH 7.5) supplemented with protease inhibitors (Roche) and 1mM DTT. The lysate was sonicated 3x for 30 seconds and incubated on ice for 30 min. Afterwards, the lysate was centrifuged at 10,000g for 10 min. 20 µl of ChromoTek GFP-Trap® agarose beads (gta, Proteintech) were applied to the supernatant and incubated overnight at 4°C on a rotating platform. The next day the sample was centrifuged at 2500g for 5 min and washed three times in IP buffer. Bound proteins were eluted with 60 µl of 1x Laemmli buffer (BioRad, 1610747) for 5 min at 95°C. IP experiments for Myc-Sipa1l3 and FLAG-Sj2bp were conducted under a slightly modified protocol. For IP experiments with Myc-Sipa1l3 0.5 µg of Sipa1l3 antibody (30544-1-AP, Proteintech) was applied to the cell lysate and incubated overnight on a rotating platform. The next day, 40 µl of prewashed protein A/G beads (Pierce) were applied to the cell lysate and incubated for 1h at 4 °C. Afterwards the beads were washed 3 times and bound proteins were eluted with 1x Laemmli buffer. IP experiments with FLAG-Sj2bp were performed with 20 µl ChromoTek DYDDDDK Fab-Trap® (ffa, Proteintech). Prior to the IP cell lysates were centrifuged at 20,000g for 30 min at 4°C. Afterwards, cell lysates were incubated with agarose overnight and washed 3 times at 4°C. Bound proteins were eluted with 60 µl of 1x Laemmli buffer.

Revised version: Immunoprecipitation experiments were carried out as previously described (47). 1.5 ×106 HEK293T cells were plated on 60mm dishes and co-transfected at the next day with 2 µg of each vector using Geneporter2 (Amsbio) transfection reagent. 24h after transfection, cells were rinsed in phosphate buffered saline (pH 7.4) and transferred to a reaction tube. The cells were centrifuged at 5000g for 5 min and afterwards homogenized in immunoprecipitation buffer (IP buffer) containing 200mM NaCl, 2m EDTA, 1% Tx-100, 50mM Tris HCL (pH 7.5) supplemented with protease inhibitors (Roche) and 1mM DTT. The lysate was sonicated 3x for 30 seconds and incubated on ice for 30 min. Afterwards, the lysate was centrifuged at 10,000g for 10 min. 20 µl of ChromoTek GFP-Trap® agarose beads (gta, Proteintech) were applied to the supernatant and incubated overnight at 4°C on a rotating platform. The next day the sample was centrifuged at 2500g for 5 min and washed three times in IP buffer. Bound proteins were eluted with 60 µl of 1x Laemmli buffer (BioRad, 1610747) for 5 min at 95°C. IP experiments for Myc-Sipa1l3 and FLAG-Sj2bp were conducted under a slightly modified protocol. For IP experiments with Myc-Sipa1l3 0.5 µg of Sipa1l3 antibody (30544-1-AP, Proteintech) was applied to the cell lysate and incubated overnight on a rotating platform. The next day, 40 µl of prewashed protein A/G beads (Pierce) were applied to the cell lysate and incubated for 1h at 4 °C. Afterwards the beads were washed 3 times and bound proteins were eluted with 1x Laemmli buffer. IP experiments with FLAG-Sj2bp were performed with 20 µl ChromoTek DYDDDDK Fab-Trap® (ffa, Proteintech). Prior to the IP cell lysates were centrifuged at 20,000g for 30 min at 4°C. Afterwards, cell lysates were incubated with agarose overnight and washed 3 times at 4°C. Bound proteins were eluted with 60 µl of 1x Laemmli buffer.

Western blot

Earlier version: Protein samples were separated via SDS-PAGE containing 7.5% (for Sipa1l3 blots) or 10% polyacrylamide. Proteins were transferred to nitrocellulose membranes using the transblot turbo system as previously described (51). Blots were blocked in 5% dried milk powder in TBST (20 mM Tris pH 7.5, 150 mM NaCl, 0.2% Tween 20) for 30 min at RT and incubated overnight in the primary antibodies diluted in blocking solution at 4°C on a rotating platform. The following primary antibodies were used: Cx36, 1:500 (Clone: 8F6, MAB3045, Millipore); Cx36, 1:500 (Clone: 1E5H5, 37-4600, Thermofisher), SIPA1L3, 1:4000 (30544-1-AP, Proteintech), GFP, 1:1000 (2956, Cell Signaling), V5 1:1000 (R960-2, Thermofisher). Biotinylated proteins detected with high sensitivity streptavidin-HRP, 1:1000 (Abcam). At the next day, blots were washed 3×10 min in TBST and incubated with the secondary antibodies, diluted in the blocking solution as described previously (51). The following secondary antibodies were used: goat anti mouse HRP, 1:1000 (34130) and goat anti rabbit, 1:1000 (34160, Thermofisher). Afterwards, the blots were washed 3×10 min in TBST and incubated with the ECL solution (32106, Thermofisher) for detection. ECL was detected with the Biorad ChemicDoc™ MP Imaging System.

Revised version: Protein samples were separated via SDS-PAGE containing 7.5% (for Sipa1l3 blots) or 10% polyacrylamide. Proteins were transferred to nitrocellulose membranes using the transblot turbo system as previously described (52). Blots were blocked in 5% dried milk powder in TBST (20 mM Tris pH 7.5, 150 mM NaCl, 0.2% Tween 20) for 30 min at RT and incubated overnight in the primary antibodies diluted in blocking solution at 4°C on a rotating platform. The following primary antibodies were used: Cx36, 1:500 (Clone: 8F6, MAB3045, Millipore); Cx36, 1:500 (Clone: 1E5H5, 37-4600, Thermofisher), SIPA1L3, 1:4000 (30544-1-AP, Proteintech), GFP, 1:1000 (2956, Cell Signaling), V5 1:1000 (R960-2, Thermofisher). Biotinylated proteins detected with high sensitivity streptavidin-HRP, 1:1000 (Abcam). At the next day, blots were washed 3×10 min in TBST and incubated with the secondary antibodies, diluted in the blocking solution as described previously (52). The following secondary antibodies were used: goat anti mouse HRP, 1:1000 (34130) and goat anti rabbit, 1:1000 (34160, Thermofisher). Afterwards, the blots were washed 3×10 min in TBST and incubated with the ECL solution (32106, Thermofisher) for detection. ECL was detected with the Biorad ChemicDoc™ MP Imaging System.

Confocal microscopy and image analysis

Earlier version: Fluorescence images were acquired with a confocal laser scanning microscope (Zeiss LSM 800) using a 60x oil objective and the Airy scan function (Pixel size: 50 nm). Confocal scans were processed using Fiji (ImageJ2 version 2.14.0/1.54f) software (69). The threshold was adjusted using the Triangle function for the entire z-stack (7 slices; 1.08 μm). Two regions of interest (ROI) with the size of 20.07 x 32.91 μm2 were positioned to cover the ON and OFF IPL layers. Frequency and size of the immunoreactive puncta were measured for each channel using the “analyze particles” function. Colocalization of ZO-1 and Sipa1l3 was quantified using the Colocalization plugin in Fiji. The number of colocalized puncta and the colocalization was measured using the Analyze particle function with a size an exclusion criterion of >1002 pixels. Following the quantification, the size, density and percentage of the colocalized puncta were averaged for each protein.

Revised version: Fluorescence images were acquired with a confocal laser scanning microscope (Zeiss LSM 800) using a 60x oil objective and the Airy scan function (Pixel size: 50 nm). Confocal scans were processed using Fiji (ImageJ2 version 2.14.0/1.54f) software (70). The threshold was adjusted using the Triangle function for the entire z-stack (7 slices; 1.08 μm). Two regions of interest (ROI) with the size of 20.07 x 32.91 μm2 were positioned to cover the ON and OFF IPL layers. Frequency and size of the immunoreactive puncta were measured for each channel using the “analyze particles” function. Colocalization of ZO-1 and Sipa1l3 was quantified using the Colocalization plugin in Fiji. The number of colocalized puncta and the colocalization was measured using the Analyze particle function with a size an exclusion criterion of >1002 pixels. Following the quantification, the size, density and percentage of the colocalized puncta were averaged for each protein.

Volume construction

Earlier version: Retinal Connectome 2 (RC2) is an ultrastructural dataset acquired from the retina of a light-adapted 5-month-old female C57BL/6J mouse (70). Methods concerning tissue acquisition and processing, volume assembly, visualization, and annotation following that extensively detailed for RC1 (71, 72). In short, the RC2 dataset spans the ganglion cell layer through outer nuclear layer of a 0.25 mm-diameter field of retina. RC2 was constructed from 1335 serial TEM sections (70 nm thick), captured at 2.18 nm/pixel via Automated Transmission Electron Microscopy at 5000X and combined with 31 optical sections intercalated throughout the volume. These optical sections were probed for small-molecule signals for computational molecular phenotyping (73). Sections were aligned into a single volume using the NCR ToolSet, which has since been replaced by Nornir (RRID:SCR_016458). All protocols were in accord with Institutional Animal Care and Use protocols of the University of Utah, the ARVO Statement for the Use of Animals in Ophthalmic and Visual Research, and the Policies on the Use of Animals and Humans in Neuroscience Research of the Society for Neuroscience.

Revised version: Retinal Connectome 2 (RC2) is an ultrastructural dataset acquired from the retina of a light-adapted 5-month-old female C57BL/6J mouse (71). Methods concerning tissue acquisition and processing, volume assembly, visualization, and annotation following that extensively detailed for RC1 (72, 73). In short, the RC2 dataset spans the ganglion cell layer through outer nuclear layer of a 0.25 mm-diameter field of retina. RC2 was constructed from 1335 serial TEM sections (70 nm thick), captured at 2.18 nm/pixel via Automated Transmission Electron Microscopy at 5000X and combined with 31 optical sections intercalated throughout the volume. These optical sections were probed for small-molecule signals for computational molecular phenotyping (74). Sections were aligned into a single volume using the NCR ToolSet, which has since been replaced by Nornir (RRID:SCR_016458). All protocols were in accord with Institutional Animal Care and Use protocols of the University of Utah, the ARVO Statement for the Use of Animals in Ophthalmic and Visual Research, and the Policies on the Use of Animals and Humans in Neuroscience Research of the Society for Neuroscience.

Statistical analysis

Data sets acquired in this study were analyzed with GraphPad Prism 8. Data are shown as mean ± SEM. Normality was tested using the Anderson–Darling and the D’Agostino-Pearson test. Significance was tested using the two-tailed Mann–Whitney U test. For multiple comparisons a one-way ANOVA was performed.